<font color ='Navy'><center> __EMship+(Advanced Design of Sustainable
Ships and Offshore Structures)__ </center></font>
<br> Offshore Structures & Digital Twin <br>  __Structural health monitoring for offshore structures__ <br> 
***
***
<font color='Navy' size=6><b><center>Working with experimental data</center></b></font> 
***
    
<div class="alert alert-block alert-info">
<center><b>Don't forget: It is important that you run each cell of the notebook. To do so select a cell (it will be highlighted) and press shift+enter on your keyboard or the play button in the menu above.</b></center></div> 

```python
# let's import both numpy and matplotlib.pyplot 
# the "as" in the code bellow is to tell python that if we reference to the package using np and plt 
import numpy as np
import matplotlib.pyplot as plt 
%matplotlib widget
```

We need also this function from the first part of the workshop.

```python
def calc_spectrum(signal, fs):
    """
    Calculates the FFT spectrum of signal sampled at the frequency fs
    
    signal - the signal of which you want to calculate the fft spectrum
    fs  - the sampling frequency in Hz
    
    returns:
    fft_freq - the frequency vector
    fft_value - the results of the 
    
    """
    from numpy.fft import fft, fftfreq, fftshift # We already imported the functions you'll need
    
    fft_values = fft(signal)/fs # 
    fft_freq = fftfreq(len(signal), 1/fs) # NOTE: 1/fs!!! 
    

    ## This additional step makes the plots look better, no need to worry about it :-)
    fft_values = fftshift(fft_values)
    fft_freq = fftshift(fft_freq)
        
    return fft_freq, fft_values
```

# 2. Offshore wind turbines
## Context

In today's session we will analyse the measurement data from an offshore wind turbine. The measurements that were shared are from two accelerometers installed in the nacelle of the wind turbine. The nacelle is the box at the top of the tower that houses the main drive shaft, gearbox and generator. We call the direction aligned with the wind direction the Fore-Aft direction (FA) and the direction cross-wind the sideways (SS) direction.

We will use the measurement to analyze the dynamic properties of the wind turbine and investigate it's susceptibility to different types of loading.  

<br><center><b>Figure 1: Two offshore wind turbines</b></center>
<img src="https://ilvo.vlaanderen.be/uploads/images/_900xAUTO_fit_center-center_none/RS2522_Thornton.jpg" width="600">


<div class="alert alert-block alert-info">
In the following exercices, we will use the same libraries as the one used in the previously. But we'll add a new package called `pandas`. `pandas` is a very useful package when it comes to importing and processing big datasets. It can handle all types of data and today we will use it to import measurement data. We'll use it later on in the session.
</div>

## Exercise 1: the rotor stop
To start the exercise we will first import some data from an actual offshore wind turbine in the North sea.
For this purpose we will use the package `pandas` and the `read_csv` command to read data from a csv file with the data stored in it. Once imported the data is stored in a so-called `Dataframe` (hence why we call it `df`). Using the `df.head()` command we show the first five lines of the Dataframe.

```python
import pandas as pd

df = pd.read_csv('offshore_wind_data_rotor_stop.csv', header=0)
df.head()
```


<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">
<b> Assignment 1.1:</b>  Use the cell below to find key information about the measurement dataset. Define:

        - The sample frequency in Hz: `f_s`
        - The duration in s: `t_e`
        - The measurement unit of the acceleration: `unit_str`
 </div>
<div class="alert alert-block alert-info">
<center><b>Tips: to learn how to access/index pandas dataframe check this  
    <a href="http://datacamp-community-prod.s3.amazonaws.com/f04456d7-8e61-482f-9cc9-da6f7f25fc9b">link. </a> <br>
   To find the duration t_e, you either have to access the last row of the pandas dataframe or you have to compute the length of the dataframe
 </b></center></div>



<!-- grade_id: cell-5fb1090bdee5e04c, points: 0, grade: false, solution: true -->
```python
# add your code here
```

Some small checks

<!-- grade_id: cell-9c094521c0dccf1f, points: 1, grade: true, solution: false -->
```python
assert f_s*t_e == len(df) # See if the total number of samples adds up
```

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

<b> Assignment 1.2:</b> use the `plt.plot` command to visualize the measurement data.
</div>


<div class="alert alert-block alert-info">
<center><b>Tips: Make your plot beautiful by adding: grid, legend, xlabel and ylabel. Google is your friend   
 </b></center></div>

<!-- grade_id: cell-a349d46d013f6dab, points: 0, grade: false, solution: true -->
```python
plt.figure(figsize=(9,5))

# add your code here

plt.show()
```

In this dataset the turbine performs a so-called **rotor stop**, during a rotor stop the turbine shuts down and the wind load is released. In this video you can see an onshore wind turbine do a rotor stop:


```python
from IPython.display import YouTubeVideo, Markdown

YouTubeVideo('_qVXkAWtH60', width=800, height=300)
```

In a way you could look at this rotor stop as an impulse response, so let us try to determine some of the wind turbine properties from it.
<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

<b> Assignment 1.3:</b> Can you isolate the data from the rotor stop in the FA and SS directions? Use `df.to_numpy` to extract data from a  dataframe into a numpy array (which is what we used in the previous sessions). Store your variable as `fa_rotor_stop` and `fss_rotor_stop`, also store the corresponding timevector (t=0s, on the rotor stop) in `t_rotor_stop`.
</div>

<!-- grade_id: cell-12abe6aa96ba6f19, points: 0, grade: false, solution: true -->
```python
# add your code here

plt.figure(figsize=(9,5))

plt.plot(t_rotor_stop, fa_rotor_stop, label='FA during rotor stop')
plt.plot(tss_rotor_stop, fss_rotor_stop, label='SS during rotor stop')
plt.xlim([t_rotor_stop[0], t_rotor_stop[-1]])
plt.xlabel('Time (s)')
plt.ylabel('Acceleration (g)')
plt.grid(which='both', linestyle=':')
plt.legend()
plt.show()
```

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

<b> Assignment 1.4:</b> Convert the impulse response data to the frequency domain using `calc_spectrum` and plot the resulting spectra. Use these spectra to estimate the first natural frequency of the structure. Utilize the `plt.semilogy()` function from matplotlib to achieve a logarithmic y-axis scale.

</div>

<!-- grade_id: cell-8bfe0eaea01a1338, points: 0, grade: false, solution: true -->
```python
plt.figure(figsize=(9,5))

# add your code here

plt.xlabel('Frequency (Hz)')
plt.ylabel('|Acceleration| (g)')
plt.grid(which='both', linestyle=':')
plt.legend()
plt.show()

```

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

<b> Assignment 1.4b:</b> What resonance frequency do you conclude for this structure? Set a variable `w_n` accordingly (best in rad/s).

</div>

<!-- grade_id: cell-cc86fb49caf7b5fe, points: 0, grade: false, solution: true -->
```python
# add your code here
```

And let us check 

<!-- grade_id: cell-b8cab3d2f5a1b925, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert (abs(w_n- 1.85) <0.25)
###END HIDDEN TESTS
```

Finally we need to talk about the damping ratio. Do you already have a rough estimate of the damping value? Is it 0.1%, 1% or 10% damping. We can try to answer this by fitting the exponential decay of the impulse response to the data.

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

<b> Assignment 2.5:</b> Using the widget below can you match the damping ratio of the wind turbine? What value do you find, for the damping? Does that seem possible?

</div>

```python
import iplot.interactive_plot_2 as ip
ip.plot_interactive()
```

<!-- grade_id: cell-9c3f367bb124b948, points: 1, grade: true, solution: true -->
<details>
    <summary><b> Click to get the answer</b></summary>
    The damping matches at around 0.24%. That is realistic for a fully steel structure.
</details> 

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

<b> Assignment 1.5b:</b> Just for the record can you store your damping value (%) in `damping_owt`?

</div>

<!-- grade_id: cell-4ecc6dc3976301e1, points: 0, grade: false, solution: true -->
```python
# add your code here
```

Let us verify you provided an answer:

<!-- grade_id: cell-bed2667cf36c347b, points: 1, grade: true, solution: false -->
```python
assert damping_owt

### BEGIN HIDDEN TESTS
assert (damping_owt>0.2) and (damping_owt<0.3)
### END HIDDEN TESTS
```

## Exercise 2: the parked turbine

In the field of offshore wind energy the behaviour of a parked turbine (i.e. when it is not producing power) is almost equally important as the behaviour during operating conditions. At some sites the fatigue life is actually determined by the amount of time a turbine spends in parked conditions. A statement like "10% of downtime accounts for 90% of the fatigue life" are not uncommon.

In this exercise we will go to a **different wind farm**, and study the behaviour a turbine that is parked. This time we will have 6 measurement signals, recorded at three different heights of the turbine. In offshore wind turbines the height is often related to a reference called LAT or Lowest Astronomical Tide, the lowest occuring tidal level. So a sensor labeled LAT015 is installed 15 meters above this reference level, and a sensor labeled LAT097 is installed 97m above this reference level. 

<img src="./Images/OWT_sensors.jpg" height=400 />

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

<b> Assignment 2.1:</b> Can you use pandas to import the measurement data that is stored in the file `offshore_wind_turbine_parked.csv`? Store the import in the variable `df`.

</div>

<!-- grade_id: cell-1ea7cab48136e418, points: 0, grade: false, solution: true -->
```python
# add your code here
```

Let us check your import

<!-- grade_id: cell-1e38d184290b30d8, points: 1, grade: true, solution: false -->
```python
assert len(df.columns) == 7, f'Did you import the new data? Your import  has {len(df.columns)} columns, we were expecting 7'
```

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

<b> Assignment 2.2:</b> Use the cell below to find key information about the measurement dataset. Define:

        - The sample frequency in Hz: `f_s`
        - The duration in s: `t_e`
        - The measurement unit of the acceleration: `unit_str`
</div>

<!-- grade_id: cell-f01dea84f18ec218, points: 0, grade: false, solution: true -->
```python
# add your code here
```

Some small checks

<!-- grade_id: cell-197d96d009cb7aed, points: 1, grade: true, solution: false -->
```python
assert f_s*t_e == len(df), f'Something does not add up, the data contains {len(df)} samples, yet t_e*f_s={t_e}*{f_s}={t_e*f_s}' # See if the total number of samples adds up
```

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

<b> Assignment 2.3:</b> use the `plt.plot` command to visualize the measurement data. Can you make sure to plot all 6 measurement signals?
</div>

<!-- grade_id: cell-c450171e7fa7e829, points: 0, grade: false, solution: true -->
```python
plt.figure(figsize=(9,5))

# add your code here

plt.show()
```

This timeseries is a bit harder to interpret than that of the previous turbine. Perhaps transforming it into the frequency domain is more helpful in interpreting the results....

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

<b> Assignment 2.4:</b> Convert the measurement data to the frequency domain using `calc_spectrum` and plot the resulting amplitude spectra. Use these spectra to estimate the first natural frequency of the structure. Zoom in on the range between 0 and 2 Hz.

</div>

<!-- grade_id: cell-e051b47ab95a298f, points: 0, grade: false, solution: true -->
```python
plt.figure(figsize=(9,5))

# add your code here

plt.xlabel('Frequency (Hz)')
plt.ylabel('|Acceleration| (g)')
plt.grid(which='both', linestyle=':')
plt.show()
```


<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

<b> Assignment 2.5:</b> Try to answer following questions below. 

- How many modes can you spot between 0 and 2Hz?
- At which frequencies are they located? 
- Do you observe differences between the different sensor levels? Can you spot all modes at every level?
- Which mode is dominating the spectrum, i.e. carries the most energy?

</div>

<!-- grade_id: cell-38d940ce2e1333a7, points: 1, grade: true, solution: true -->
<details>
    <summary><b> Click to get the answer</b></summary>
    
- There are 3 notable modes visible in the spectrum. Potentially there are more hidden underneath the noise.

- The first is around 0.23Hz, the second at around 0.75Hz and the third at around 1.34Hz
- Depending on which mode you consider a different sensor has the highest energy. E.g. for the first mode the sensors that carry the most energy are the top (LAT097) sensors. Meanwhile for the mode at 1.34Hz the lower sensors (LAT069 and LAT015) carry more energy
- The spectrum is dominated by the first mode at 0.23Hz, it contains by far the most energy.  
</details> 

These measurements can also provide us some insight into the mode shapes. While we can't derive the modeshapes directly from the spectrum we can visualize the so-called Operational Deflection Shapes (ODS) from it. These show how the structure vibrates at a given frequency. In particular at (well-seperated) resonance frequencies, these will resemble the mode shapes.

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

<b> Assignment 2.6:</b> Use the widget below to visualize the ODS at various frequencies for the FA signals. Can you relate the observed ODS to expected structural modes of the system?

</div>

```python
import iplot.interactive_plot_3 as ip
ip.plot_interactive()
```

Addtionally, I would like to share with you a very usefull tip when handling real-world data. In this sessions we talke about calculating spectra and windows. We did not speak of noise. The latter is always present and to some level will limit what we can do with measurements. When noise is too bad our signal will get flooded by noise and we won't be able to see what is going on. To resolve this we can work with averaging. A proper way of doing this is the so-called Welch's method.

Welch's method is a way to compute the **Power Spectral Density (PSD)** of a signal. Welch's method is available from `scipy.signal` as `welch`. `welch` does not only provide an implementation of Welch's method, it also simplifies the process of using a window.

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

<b> Assignment 2.7:</b> Import `welch` from `scipy.signal` and use it to calculate the PSD of the measurement data at LAT015 in the FA-direction contained in the dataframe `df`. Store the resulting frequency vectors as `f_psd` and the PSD itself as `psd`. Consider different settings of the windowing length and the window type. Also set the variable `psd_y_label` label to correct engineering units knowing the time domain signal is `g`. e.g. `psd_y_label='g'`
</div>


<!-- grade_id: cell-390357e805648d13, points: 0, grade: false, solution: true -->
```python
# add your code here
```

Let us have a look at your results:

<!-- grade_id: cell-4c69b3fdc57e6b9b, points: 1, grade: true, solution: false -->
```python
# add your code here
```

<div class="alert alert-block alert-info">
<center><b>This concludes this part of the workshop, hope you learned a thing or two..</b></center></div> 